---
title: "Tutorial 02"
format: live-html
engine: knitr
# prevent any project filters from interfering
sidebar:
  style: docked
  contents:
    - text: "Q1 — Sum to 10"
      href: "#q1"
    - text: "Q2 — Division"
      href: "#q2"
    - text: "Q3 — Histogram"
      href: "#q3"
---

{{< include ../_extensions/r-wasm/live/_knitr.qmd >}} {{< include ../_extensions/r-wasm/live/_gradethis.qmd >}}

## Q1 — Quick five-number summary {#q1}

Create mpg_summary as the output of summary(mtcars$mpg).

```{webr}
#| exercise: ex_summary_mpg
#| exercise.lines: 2
#| echo: false
mpg_summary <- NULL
```

::: {.hint exercise="ex_summary_mpg"}
Use summary(mtcars$mpg).
:::

::: {.solution exercise="ex_summary_mpg"}
```{webr}
#| exercise: ex_summary_mpg
#| solution: true
mpg_summary <- summary(mtcars$mpg)
```
:::

```{webr}
#| exercise: ex_summary_mpg
#| check: true
gradethis::grade_this({
  e <- get0(".envir_result", ifnotfound = parent.frame())
  if (!exists("mpg_summary", envir = e)) fail("Create `mpg_summary`.")
  s <- get("mpg_summary", envir = e)
  if (!is.numeric(s) || length(s) < 6) fail("`summary(mtcars$mpg)` returns a named numeric vector (Min., 1st Qu., Median, Mean, 3rd Qu., Max.).")
  need <- c("Min.", "1st Qu.", "Median", "Mean", "3rd Qu.", "Max.")
  if (!all(need %in% names(s))) fail("`mpg_summary` should come from `summary(...)` with standard names.")
  pass("✅ Summary looks good.")
})
```

## Q2 — Mean, variance, and sd {#q2}

Create a named numeric vector stats_wt containing the mean, var, and sd of mtcars$wt with na.rm = TRUE.

```{webr}
#| exercise: ex_stats_wt
#| exercise.lines: 3
#| echo: false
stats_wt <- NULL
```

::: {.hint exercise="ex_stats_wt"}
c(mean = mean(...), var = var(...), sd = sd(...))
:::

::: {.solution exercise="ex_stats_wt"}
```{webr}
#| exercise: ex_stats_wt
#| solution: true
stats_wt <- c(
  mean = mean(mtcars$wt, na.rm = TRUE),
  var  = var(mtcars$wt,  na.rm = TRUE),
  sd   = sd(mtcars$wt,   na.rm = TRUE)
)
```
:::

```{webr}
#| exercise: ex_stats_wt
#| check: true
gradethis::grade_this({
  e <- get0(".envir_result", ifnotfound = parent.frame())
  if (!exists("stats_wt", envir = e)) fail("Define `stats_wt`.")
  v <- get("stats_wt", envir = e)
  if (!is.numeric(v) || is.null(names(v))) fail("Return a **named numeric vector**.")
  need <- c("mean","var","sd")
  if (!setequal(names(v), need)) fail("Names must be exactly: mean, var, sd.")
  # target values (computed once here)
  tgt <- c(mean = mean(mtcars$wt), var = var(mtcars$wt), sd = sd(mtcars$wt))
  ok <- all(mapply(function(x,y) isTRUE(all.equal(x,y,tol=1e-8)), v[need], tgt[need]))
  if (ok) pass("✅ Correct mean/var/sd.")
  else fail("Recheck the formulas and use `na.rm = TRUE` (even though there are no NAs here).")
})
```

## Q3 — Z-scores

Compute z-scores for mtcars$mpg and store them in z_mpg.

::: {.callout-note title="Info"}
Remember: $$z = \frac{x - \mu }{\sigma/\sqrt{n})} $$
:::


```{webr}
#| exercise: ex_z_mpg
#| exercise.lines: 3
#| echo: false
z_mpg <- NULL

```

::: {.hint exercise="ex_z_mpg"}
(x - mean(x, na.rm=TRUE)) / sd(x, na.rm=TRUE)
:::

::: {.solution exercise="ex_z_mpg"}
```{webr}
#| exercise: ex_z_mpg
#| solution: true
x <- mtcars$mpg
z_mpg <- (x - mean(x, na.rm = TRUE)) / sd(x, na.rm = TRUE)

```
:::

```{webr}
#| exercise: ex_z_mpg
#| check: true
gradethis::grade_this({
  e <- get0(".envir_result", ifnotfound = parent.frame())
  if (!exists("z_mpg", envir = e)) fail("Create `z_mpg`.")
  z <- get("z_mpg", envir = e)
  if (!is.numeric(z) || length(z) != nrow(mtcars)) fail("`z_mpg` should be a numeric vector of length nrow(mtcars).")
  ok_mean <- abs(mean(z)) < 1e-8
  ok_sd   <- isTRUE(all.equal(sd(z), 1, tol = 1e-8))
  if (ok_mean && ok_sd) pass("✅ Looks like proper z-scores (mean ≈ 0, sd ≈ 1).")
  else fail("Compute z = (x - mean(x))/sd(x) with `na.rm=TRUE`.")
})
```

## Q4 — summary() on the whole data frame 

Create mtcars_summary as summary(mtcars).

```{webr}
#| exercise: ex_summary_df
#| exercise.lines: 2
#| echo: false
mtcars_summary <- NULL

```

::: {.hint exercise="ex_summary_df"}
Use summary(mtcars).
:::

::: {.solution exercise="ex_summary_df"}
```{webr}
#| exercise: ex_summary_df
#| solution: true
mtcars_summary <- summary(mtcars)
```
:::

```{webr}
#| exercise: ex_summary_df
#| check: true
gradethis::grade_this({
  e <- get0(".envir_result", ifnotfound = parent.frame())
  if (!exists("mtcars_summary", envir = e)) fail("Create `mtcars_summary`.")
  s <- get("mtcars_summary", envir = e)
  if (!is.list(s) && !is.matrix(s)) fail("`summary(mtcars)` returns a structured summary per column.")
  pass("✅ Data-frame summary created.")
})
```


## Q5 — Quick data health check with skimr 

Use skimr::skim(mtcars) and save to skim_out.

```{webr}
#| exercise: ex_skim
#| exercise.lines: 2
#| echo: false
skim_out <- NULL

```

::: {.hint exercise="ex_skim"}
skimr::skim(mtcars)
:::

::: {.solution exercise="ex_skim"}
```{webr}
#| exercise: ex_skim
#| solution: true
skim_out <- skimr::skim(mtcars)
```
:::

```{webr}
#| exercise: ex_skim
#| check: true
gradethis::grade_this({
  if (!requireNamespace("skimr", quietly = TRUE)) fail("Please install/load **skimr**.")
  e <- get0(".envir_result", ifnotfound = parent.frame())
  if (!exists("skim_out", envir = e)) fail("Store the result in `skim_out`.")
  so <- get("skim_out", envir = e)
  ok <- is.data.frame(so) && all(c("skim_type","skim_variable") %in% names(so))
  if (!ok) fail("`skim_out` should be a skim data frame.")
  pass("✅ skim summary captured.")
})
```

## Q6 — Grouped summary with dplyr (Intro to tidyverse)

Create a tibble by_cyl with one row per cyl containing:

n (row count),

mpg_mean,

hp_sd.

```{webr}
#| exercise: ex_dplyr_group
#| exercise.lines: 7
#| echo: false
by_cyl <- NULL

```

::: {.solution exercise="ex_dplyr_group"}
```{webr}
#| exercise: ex_dplyr_group
#| solution: true
by_cyl <- mtcars |>
  dplyr::as_tibble() |>
  dplyr::group_by(cyl) |>
  dplyr::summarise(
    n = dplyr::n(),
    mpg_mean = mean(mpg, na.rm = TRUE),
    hp_sd = sd(hp, na.rm = TRUE),
    .groups = "drop"
  )

```
:::

```{webr}
#| exercise: ex_dplyr_group
#| check: true
gradethis::grade_this({
  if (!requireNamespace("dplyr", quietly = TRUE)) fail("Load **dplyr**.")
  e <- get0(".envir_result", ifnotfound = parent.frame())
  if (!exists("by_cyl", envir = e)) fail("Create `by_cyl`.")
  df <- get("by_cyl", envir = e)
  need <- c("cyl","n","mpg_mean","hp_sd")
  if (!all(need %in% names(df))) fail("Include columns: cyl, n, mpg_mean, hp_sd.")
  if (nrow(df) != length(unique(mtcars$cyl))) fail("One row per `cyl`.")
  pass("✅ Grouped summary OK.")
})
```

## Q7 — Multiple summaries with across() (tidyverse)

Make by_am with one row per am and columns:

n,

mpg_mean, mpg_sd,

wt_mean, wt_sd.

```{webr}
#| exercise: ex_across
#| exercise.lines: 10
#| echo: false
by_am <- NULL
```

::: {.solution exercise="ex_across"}
```{webr}
#| exercise: ex_across
#| solution: true
by_am <- mtcars |>
  dplyr::as_tibble() |>
  dplyr::group_by(am) |>
  dplyr::summarise(
    n = dplyr::n(),
    dplyr::across(
      c(mpg, wt),
      list(mean = ~mean(.x, na.rm = TRUE), sd = ~sd(.x, na.rm = TRUE)),
      .names = "{.col}_{.fn}"
    ),
    .groups = "drop"
  )
```
:::

```{webr}
#| exercise: ex_across
#| check: true
gradethis::grade_this({
  if (!requireNamespace("dplyr", quietly = TRUE)) fail("Load **dplyr**.")
  e <- get0(".envir_result", ifnotfound = parent.frame())
  if (!exists("by_am", envir = e)) fail("Create `by_am`.")
  df <- get("by_am", envir = e)
  need <- c("am","n","mpg_mean","mpg_sd","wt_mean","wt_sd")
  if (!all(need %in% names(df))) fail("Include columns: am, n, mpg_mean, mpg_sd, wt_mean, wt_sd.")
  pass("✅ Multiple summaries computed.")
})
```

## Q8 — Histogram of mpg (ggplot2)

Make by_am with one row per am and columns:

n,

mpg_mean, mpg_sd,

wt_mean, wt_sd.

```{webr}
#| exercise: ex_across
#| exercise.lines: 10
#| echo: false
by_am <- NULL
```

::: {.solution exercise="ex_across"}
```{webr}
#| exercise: ex_across
#| solution: true
by_am <- mtcars |>
  dplyr::as_tibble() |>
  dplyr::group_by(am) |>
  dplyr::summarise(
    n = dplyr::n(),
    dplyr::across(
      c(mpg, wt),
      list(mean = ~mean(.x, na.rm = TRUE), sd = ~sd(.x, na.rm = TRUE)),
      .names = "{.col}_{.fn}"
    ),
    .groups = "drop"
  )
```
:::

```{webr}
#| exercise: ex_across
#| check: true
gradethis::grade_this({
  if (!requireNamespace("dplyr", quietly = TRUE)) fail("Load **dplyr**.")
  e <- get0(".envir_result", ifnotfound = parent.frame())
  if (!exists("by_am", envir = e)) fail("Create `by_am`.")
  df <- get("by_am", envir = e)
  need <- c("am","n","mpg_mean","mpg_sd","wt_mean","wt_sd")
  if (!all(need %in% names(df))) fail("Include columns: am, n, mpg_mean, mpg_sd, wt_mean, wt_sd.")
  pass("✅ Multiple summaries computed.")
})
```
