---
title: "Tutorial 10"
format: live-html
engine: knitr
resources:
  - life_expectancy.csv
---

{{< include ../_extensions/r-wasm/live/_knitr.qmd >}}
{{< include ../_extensions/r-wasm/live/_gradethis.qmd >}}

## Q1 - Fit the Regression Model 

For this tutorial, we will be using a Life Expectancy dataset from the World Health Organization (WHO). From your tutorial 9 knowledge, fit a simple linear regression model with Life Expectancy as X and Schooling as Y. Return the estimated slope. 

::: {.callout-note title="Preview"}
Feel free to run this code block to visualize the data.
```{webr}
#| echo: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))

le <- df[, c("Life expectancy", "Schooling")]
names(le) <- c("life_exp", "schooling")
le <- na.omit(le)

head(le, 8)


plot(le$schooling, le$life_exp, 
xlab = "Schooling (years)", ylab = "Life expectancy",
main = "Life expectancy vs Schooling")

```
:::

```{webr}
#| exercise: t10_q1_fit_slope
#| exercise.lines: 8
#| echo: false
df <- read.csv("life_expectancy.csv", check.names=FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling") #use these column names 
le <- na.omit(le)

model <- ___ ( ___ ~ ___ , data = le)
coef(model)[___]

```

:::{.hint exercise="t10_q1_fit_slope"}

Fill coef(model)[___] with right column name.

:::

::: {.solution exercise="t10_q1_fit_slope"}

```{webr}

#| exercise: t10_q1_fit_slope
#| solution: true
df <- read.csv("life_expectancy.csv", check.names=FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)

model <- lm(life_exp ~ schooling, data = le)
coef(model)["schooling"]
```
:::

```{webr}
#| exercise: t10_q1_fit_slope
#| check: true
gradethis::grade_this({
df <- read.csv("life_expectancy.csv", check.names=FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
mod <- lm(life_exp ~ schooling, data = le)
exp <- coef(mod)["schooling"]

x <- .result
ok <- is.numeric(x) && length(x) == 1L && is.finite(x)
if (!ok) fail("Return the slope as a single numeric value.")
else if (abs(x - exp) < 1e-10) pass("Correct slope.")
else fail("Slope mismatch. Check your lm() formula and coefficient name.")
})

```

## Q2 — Interpret the slope: 

Compute the predicted change in life expectancy if schooling is increased in all countries by 2 years. Return a single number.


<figure style="text-align:center">
  <img src="./images/marvel.jpg" style="width:40%">
  <figcaption>Photo by <a href="https://unsplash.com/@introspectivedsgn?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Erik Mclean</a> on <a href="https://unsplash.com/photos/a-pile-of-comics-sitting-next-to-each-other-8SeJUmfahu0?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Unsplash</a>
      </figcaption>
</figure>

::: {.callout-note title="Info"}
If the slope is $b_1$, then increasing schooling by 2 years will also double $b_1$. 
:::


::: {.callout-note title="Preview"}
Feel free to run this code block to visualize the data.
```{webr}
#| echo: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))

le <- df[, c("Life expectancy", "Schooling")]
names(le) <- c("life_exp", "schooling")
le <- na.omit(le)

head(le, 8)


plot(le$schooling, le$life_exp, 
xlab = "Schooling (years)", ylab = "Life expectancy",
main = "Life expectancy vs Schooling")

```
:::

```{webr}
#| exercise: t10_q2_two_year_effect
#| exercise.lines: 8
#| echo: false
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
model <- lm(life_exp ~ schooling, data = le)

b1 <- ___
 #transform slope here
```

::: {.hint exercise="t10_q2_two_year_effect"}
Slope will be doubled! Discuss why.
:::

::: {.solution exercise="t10_q2_two_year_effect"}

```{webr}
#| exercise: t10_q2_two_year_effect
#| solution: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
model <- lm(life_exp ~ schooling, data = le)

b1 <- coef(model)["schooling"]
2 * b1

```

:::

```{webr}
#| exercise: t10_q2_two_year_effect
#| check: true
gradethis::grade_this({
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
mod <- lm(life_exp ~ schooling, data = le)
exp <- 2 * coef(mod)["schooling"]

x <- .result
ok <- is.numeric(x) && length(x) == 1L && is.finite(x)
if (!ok) fail("Return a single numeric value (the +2 years effect).")
else if (abs(x - exp) < 1e-10) pass("Correct.")
else fail("Mismatch. Check b1 and multiply by 2.")
})

```

## Q3 — Compute $R^2$

Return the model's $R^2$


<figure style="text-align:center">
  <img src="./images/captain_vs_iron.jpg" style="width:40%">
  <figcaption>Photo by <a href="https://unsplash.com/@wacalke?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Mateusz Wacławek</a> on <a href="https://unsplash.com/photos/red-and-blue-robot-toy-t2b2svMf8ek?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Unsplash</a>
      </figcaption>
</figure>


::: {.callout-note title="Info"}
$R^2$ is the proportion of variability in life expectancy explained by schooling. 
:::

::: {.callout-note title="Preview"}
Feel free to run this code block to visualize the data.
```{webr}
#| echo: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))

le <- df[, c("Life expectancy", "Schooling")]
names(le) <- c("life_exp", "schooling")
le <- na.omit(le)

head(le, 8)


plot(le$schooling, le$life_exp, 
xlab = "Schooling (years)", ylab = "Life expectancy",
main = "Life expectancy vs Schooling")

```
:::

```{webr}
#| exercise: t10_q3_r2
#| exercise.lines: 7
#| echo: false
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
model <- lm(life_exp ~ schooling, data = le)

___(model)$___

```

::: {.hint exercise="t10_q3_r2"}
Produce model summary and return the appropriate statistic from it. 
:::

::: {.solution exercise="t10_q3_r2"}

```{webr}
#| exercise: t10_q3_r2
#| solution: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
model <- lm(life_exp ~ schooling, data = le)

summary(model)$r.squared

```
:::

```{webr}
#| exercise: t10_q3_r2
#| check: true
gradethis::grade_this({
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
mod <- lm(life_exp ~ schooling, data = le)
exp <- summary(mod)$r.squared

x <- .result
ok <- is.numeric(x) && length(x) == 1L && is.finite(x)
if (!ok) fail("Return a single numeric R-squared value.")
else if (abs(x - exp) < 1e-12) pass("Correct R-squared.")
else fail("R-squared mismatch.")
})

```

## Q4 — Computing P-value

Compute the p-value for $b_1$ for testing $H_0: b_1 = 0$. 

<figure style="text-align:center">
  <img src="./images/thor.jpg" style="width:40%">
  <figcaption>Photo by <a href="https://unsplash.com/@wacalke?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Mateusz Wacławek</a> on <a href="https://unsplash.com/photos/red-and-blue-robot-toy-t2b2svMf8ek?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Unsplash</a>
      </figcaption>
</figure>

::: {.callout-note title="Info"}
This tests whether there is evidence of a linear relationship between schooling and life expectancy.
:::

::: {.callout-note title="Preview"}
Feel free to run this code block to visualize the data.
```{webr}
#| echo: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))

le <- df[, c("Life expectancy", "Schooling")]
names(le) <- c("life_exp", "schooling")
le <- na.omit(le)

head(le, 8)


plot(le$schooling, le$life_exp, 
xlab = "Schooling (years)", ylab = "Life expectancy",
main = "Life expectancy vs Schooling")

```
:::

```{webr}
#| exercise: t10_q4_slope_pval
#| exercise.lines: 9
#| echo: false
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
mod <- lm(life_exp ~ schooling, data = le)

summary(mod)$coefficients[___, ___]

```

::: {.hint exercise="t10_q4_slope_pval"}
The second argument in coefficients represents the column 'Pr(>|t|)'. 
:::

::: {.solution exercise="t10_q4_slope_pval"}

```{webr}
#| exercise: t10_q4_slope_pval
#| solution: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
mod <- lm(life_exp ~ schooling, data = le)

summary(mod)$coefficients["schooling", "Pr(>|t|)"]

```
:::

```{webr}
#| exercise: t10_q4_slope_pval
#| check: true
gradethis::grade_this({
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
mod <- lm(life_exp ~ schooling, data = le)
exp <- summary(mod)$coefficients["schooling", "Pr(>|t|)"]

x <- .result
ok <- is.numeric(x) && length(x) == 1L && is.finite(x)
if (!ok) fail("Return a single numeric p-value.")
else if (abs(x - exp) < 1e-12) pass("Correct p-value. A really low or 0 p-value suggests that it is highly unlikely for the slope to be 0. Can you see that from the graph as well?")
else fail("P-value mismatch.")
})

```


## Q5 — 95% CI for the slope

Return a 95% confidence interval for the slope using the R built-in command confint(). 

<figure style="text-align:center">
  <img src="./images/pokeball.jpg" style="width:40%">
  <figcaption>Photo by <a href="https://unsplash.com/@thimo?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Thimo Pedersen</a> on <a href="https://unsplash.com/photos/red-and-white-ladybug-toy-on-white-and-yellow-book-dip9IIwUK6w?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Unsplash</a>
      </figcaption>
</figure>

::: {.callout-note title="Info"}
The confint() function in R is used to compute confidence intervals for one or more parameters in a fitted model. 
:::
::: {.callout-note title="Preview"}
Feel free to run this code block to visualize the data.
```{webr}
#| echo: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))

le <- df[, c("Life expectancy", "Schooling")]
names(le) <- c("life_exp", "schooling")
le <- na.omit(le)

head(le, 8)


plot(le$schooling, le$life_exp, 
xlab = "Schooling (years)", ylab = "Life expectancy",
main = "Life expectancy vs Schooling")

```
:::

```{webr}
#| exercise: t10_q5_ci_slope
#| exercise.lines: 9
#| echo: false
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
model <- lm(life_exp ~ schooling, data = le)

as.numeric(confint(___, ___))


```

::: {.hint exercise="t10_q5_ci_slope"}
Enter the right column names and research how confint() is used here. 
:::

::: {.solution exercise="t10_q5_ci_slope"}
```{webr}
#| exercise: t10_q5_ci_slope
#| solution: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
mod <- lm(life_exp ~ schooling, data = le)

as.numeric(confint(mod, "schooling"))


```
:::


```{webr}
#| exercise: t10_q5_ci_slope
#| check: true
gradethis::grade_this({
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
mod <- lm(life_exp ~ schooling, data = le)
exp <- as.numeric(confint(mod, "schooling"))

x <- .result
ok <- is.numeric(x) && length(x) == 2L && all(is.finite(x))
if (!ok) fail("Return a numeric vector of length 2: c(lower, upper).")
else if (max(abs(x - exp)) < 1e-10) pass("Correct CI.")
else fail("CI mismatch.")
})


```

## Q6 — Predict the mean life expectancy 

Predict the mean life expectancy at schooling = 12 years.



<figure style="text-align:center">
  <img src="./images/pokeball.jpg" style="width:40%">
  <figcaption>Photo by <a href="https://unsplash.com/@thimo?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Thimo Pedersen</a> on <a href="https://unsplash.com/photos/red-and-white-ladybug-toy-on-white-and-yellow-book-dip9IIwUK6w?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Unsplash</a>
      </figcaption>
</figure>

::: {.callout-note title="Info"}
A point prediction is essentially $\hat{y}$ at a certain x value. In this case, x=12.
:::

::: {.callout-note title="Preview"}
Feel free to run this code block to visualize the data.
```{webr}
#| echo: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))

le <- df[, c("Life expectancy", "Schooling")]
names(le) <- c("life_exp", "schooling")
le <- na.omit(le)

head(le, 8)


plot(le$schooling, le$life_exp, 
xlab = "Schooling (years)", ylab = "Life expectancy",
main = "Life expectancy vs Schooling")

```
:::


```{webr}

#| exercise: t10_q6_predict_12
#| exercise.lines: 8
#| echo: false
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
model <- lm(life_exp ~ schooling, data = le)

predict(___, newdata = data.frame(schooling = ___)) #enter schooling value

```

::: {.hint exercise="t10_q6_predict_12"}
Look up how to use predict(). 
:::

::: {.solution exercise="t10_q6_predict_12"}
```{webr}
#| exercise: t10_q6_predict_12
#| solution: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
model <- lm(life_exp ~ schooling, data = le)

predict(model, newdata = data.frame(schooling = 12))

```
:::

```{webr}
#| exercise: t10_q6_predict_12
#| check: true
gradethis::grade_this({
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
mod <- lm(life_exp ~ schooling, data = le)
exp <- as.numeric(predict(mod, newdata = data.frame(schooling = 12)))

x <- .result
ok <- is.numeric(x) && length(x) == 1L && is.finite(x)
if (!ok) fail("Return a single numeric prediction.")
else if (abs(x - exp) < 1e-10) pass("Correct prediction.")
else fail("Prediction mismatch.")
})


```



## Q7 — Residuals sum

Confirm if the residuals in the model sum to 0 or not. Round the sum. 


<figure style="text-align:center">
  <img src="./images/pokemons.jpg" style="width:40%">
  <figcaption>Photo by <a href="https://unsplash.com/@lifetimevalues?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Life Time Values</a> on <a href="https://unsplash.com/photos/six-pokemon-figurines-lined-up-on-a-wooden-surface-caF_z0aTPtQ?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Unsplash</a>
      </figcaption>
</figure>

::: {.callout-note title="Info"}
Usually with an intercept in the model, the residuals sum up to 0 (with a tiny numerical error).
:::

::: {.callout-note title="Preview"}
```{webr}
#| echo: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))

le <- df[, c("Life expectancy", "Schooling")]
names(le) <- c("life_exp", "schooling")
le <- na.omit(le)

mod <- lm(life_exp ~ schooling, data = le)

# Key idea: residuals balance around 0 when there's an intercept

hist(resid(mod),
main = "Histogram of residuals",
xlab = "Residuals")

boxplot(resid(mod), horizontal = TRUE,
main = "Residuals boxplot")

# A visual “sum to zero” intuition: mean residual should be ~ 0 
mean(resid(mod))

```
:::

```{webr}
#| exercise: t10_q7_resid_sum
#| exercise.lines: 8
#| echo: false
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
model <- lm(life_exp ~ schooling, data = le)

round(sum(___(___)), 6)

```

::: {.hint exercise="t10_q7_resid_sum"}
Use resid() to get residuals. 
:::

::: {.solution exercise="t10_q7_resid_sum"}
```{webr}
#| exercise: t10_q7_resid_sum
#| solution: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
model <- lm(life_exp ~ schooling, data = le)

round(sum(resid(model)), 6)


```
:::

```{webr}
#| exercise: t10_q7_resid_sum
#| check: true
gradethis::grade_this({
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
mod <- lm(life_exp ~ schooling, data = le)
exp <- round(sum(resid(mod)), 6)

x <- .result
ok <- is.numeric(x) && length(x) == 1L && is.finite(x)
if (!ok) fail("Return one rounded number.")
else if (abs(x - exp) < 1e-12) pass("Correct.")
else fail("Mismatch.")
})

```


## Q8 — Standardized residuals

Compute standardized residuals and figure out how many satisfy the condition |$r_i$| > 2.  


<figure style="text-align:center">
  <img src="./images/pokemons.jpg" style="width:40%">
  <figcaption>Photo by <a href="https://unsplash.com/@lifetimevalues?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Life Time Values</a> on <a href="https://unsplash.com/photos/six-pokemon-figurines-lined-up-on-a-wooden-surface-caF_z0aTPtQ?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Unsplash</a>
      </figcaption>
</figure>

::: {.callout-note title="Info"}
Standardized residuals are useful for spotting outliers in data. If the condition |$r_i$| > 2 is satisfied, it could be a common red flag. 
:::

::: {.callout-note title="Preview"}
```{webr}
#| echo: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))

le <- df[, c("Life expectancy", "Schooling")]
names(le) <- c("life_exp", "schooling")
le <- na.omit(le)

mod <- lm(life_exp ~ schooling, data = le)

r <- rstandard(mod)

# Show the distribution and typical range without counting > 2

hist(r,
main = "Standardized residuals",
xlab = "rstandard(mod)")
abline(v = c(-2, 2), lwd = 2)

# Show the top few absolute values (but don't reveal the count)

head(sort(abs(r), decreasing = TRUE), 10)


```
:::

```{webr}
#| exercise: t10_q8_std_resid_count
#| exercise.lines: 9
#| echo: false
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
model <- lm(life_exp ~ schooling, data = le)

r <- ___(model)
sum(abs(r) > ___)

```

::: {.hint exercise="t10_q8_std_resid_count"}
Use rstandard() to get standardized residuals. 
:::

::: {.solution exercise="t10_q8_std_resid_count"}
```{webr}
#| exercise: t10_q8_std_resid_count
#| solution: true
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
model <- lm(life_exp ~ schooling, data = le)

r <- rstandard(model)
sum(abs(r) > 2)

```
:::

```{webr}
#| exercise: t10_q8_std_resid_count
#| check: true
gradethis::grade_this({
df <- read.csv("life_expectancy.csv", check.names = FALSE)
names(df) <- trimws(names(df))
le <- df[, c("Life expectancy","Schooling")]
names(le) <- c("life_exp","schooling")
le <- na.omit(le)
mod <- lm(life_exp ~ schooling, data = le)
exp <- sum(abs(rstandard(mod)) > 2)

x <- .result
ok <- is.numeric(x) && length(x) == 1L && is.finite(x)
if (!ok) fail("Return a single integer count.")
else if (abs(x - exp) < 1e-12) pass("Correct.")
else fail("Count mismatch.")
})


```







