---
title: "Tutorial 04: One Sample Confidence Intervals"
format: live-html
engine: knitr
---

{{< include ../_extensions/r-wasm/live/_knitr.qmd >}}
{{< include ../_extensions/r-wasm/live/_gradethis.qmd >}}


## Q1 — Z Critical Value

Print the 95% 2-sided $Z^*$ critical value when $\alpha = 0.05$. 

::: {.callout-note title="Info"}
Remember: Critical values represent points on a distribution which play an important role in both confidence intervals and hypothesis testing. They are the cutoffs on a reference distribution that set the width of confidence intervals.
:::


```{webr}
#| exercise: q1_zstar_95
#| exercise.lines: 1
#| echo: false

```

::: {.hint exercise="q1_zstar_95"}
Use qnorm.
:::

::: {.solution exercise="q1_zstar_95"}
```{webr}
#| exercise: q1_zstar_95
#| solution: true
qnorm(0.975)

```
:::

```{webr}
#| exercise: q1_zstar_95
#| check: true
gradethis::grade_this({
exp <- qnorm(0.975)
x <- .result
if (!is.numeric(x) || length(x)!=1L || !is.finite(x)) fail("Print a single numeric value.")
if (abs(x - exp) < 1e-6) pass("✅ Correct Z* for 95%.")
else fail("Not quite — check the percentile for a two-sided 95% interval.")
})

```

## Q2 — One Sample Z Confidence Interval 

The penguins dataset contains measurements for 3 penguin species from the Palmer Archipelago. Use body_mass_g (grams). Assume the population SD is known: σ = 450 g. Construct a 90% Z-interval for the true mean body mass. Print c(lower, upper).

::: {.callout-note title="Info"}
When σ is known, the margin of error uses the standard normal critical value.
:::

::: {.callout-note title="Preview"}
Run this code chunk to get a glimpse of the dataset. Feel free to change the values to visualize more/less number of rows. 
```{webr}
#| echo: true
if (!requireNamespace("palmerpenguins", quietly = TRUE)) webr::install("palmerpenguins")
library(palmerpenguins)
df <- na.omit(palmerpenguins::penguins[, c("species","island","sex","body_mass_g")])
head(df, 10)
summary(df$body_mass_g)

```
:::
::: {.callout-note title="Info"}
When σ is known,
$$
\mathrm{CI}_{1-\alpha}:\ \bar{x} \pm x_{1-\alpha/2}\,\frac{\sigma}{\sqrt{n}}.
$$
:::

```{webr}
#| exercise: q2_z_ci_penguins
#| exercise.lines: 4
#| echo: false
if (!requireNamespace("palmerpenguins", quietly = TRUE)) webr::install("palmerpenguins")
library(palmerpenguins)
x <- na.omit(palmerpenguins::penguins$body_mass_g)
xbar <-
n <-
z <-
sigma <- 450




```

::: {.hint exercise="q2_z_ci_penguins"}
Compute the sample mean, pair with the known σ and your Z critical value. Output c(lower, upper).
:::

::: {.solution exercise="q2_z_ci_penguins"}
```{webr}
#| exercise: q2_z_ci_penguins
#| solution: true
if (!requireNamespace("palmerpenguins", quietly = TRUE)) webr::install("palmerpenguins")
library(palmerpenguins)
x <- na.omit(palmerpenguins::penguins$body_mass_g)
xbar <- mean(x)
n <- length(x)
z <- qnorm(0.95) # 90% two-sided
sigma <- 450
xbar + c(-1,1) * z * sigma / sqrt(n)
```
:::

```{webr}
#| exercise: q2_z_ci_penguins
#| check: true
gradethis::grade_this({
if (!requireNamespace("palmerpenguins", quietly = TRUE)) webr::install("palmerpenguins")
library(palmerpenguins)
x <- na.omit(palmerpenguins::penguins$body_mass_g)
xbar <- mean(x); n <- length(x); z <- qnorm(0.95); sigma <- 450
exp <- xbar + c(-1,1) * z * sigma / sqrt(n)

xres <- .result
ok <- is.numeric(xres) && length(xres)==2L && all(is.finite(xres))
if (!ok) fail("Print a numeric vector c(lower, upper).")
if (max(abs(xres - exp)) < 1e-6) pass("✅ Correct 90% Z-interval for penguin body mass.")
else fail("Recheck mean, n, z*, and the known σ.")
})


```


## Q3 — T Critical Value

PlantGrowth is a dataset that contains results from an experiment to compare yields (as measured by dried weight of plants) obtained under a control and two different treatment conditions. Print the 95% two-sided t* using PlantGrowth's weight column when $\alpha = 0.05$. 


<figure style="text-align:center">
  <img src="./images/plant_growth.jpg" style="width:60%">
  <figcaption>Photo by <a href="https://unsplash.com/@arnold_nagy?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Nagy Arnold</a> on <a href="https://unsplash.com/photos/green-leafed-bonsai-tree-against-black-background-X_IvVDuHvDQ?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Unsplash</a>
      </figcaption>
</figure>

::: {.callout-note title="Preview"}
Run this code chunk to get a glimpse of the dataset. Feel free to change the values to visualize more/less number of rows. 
```{webr}
#| echo: true
plantgrowth_data <- datasets::PlantGrowth
head(plantgrowth_data, 10)
```
:::

```{webr}
#| exercise: q3_tstar_95
#| exercise.lines: 1
#| echo: false
data("PlantGrowth")
n <- 

```

::: {.hint exercise="q3_tstar_95"}
Find n from the weights vector and print the two-sided 95% t critical value using df = n − 1. Refer to the previous tutorial for a hint about the R command used to print t quantiles.
:::

::: {.solution exercise="q3_tstar_95"}
```{webr}
#| exercise: q3_tstar_95
#| solution: true
data("PlantGrowth")
n <- length(PlantGrowth$weight)
qt(0.975, df = n - 1)

```
:::

```{webr}
#| exercise: q3_tstar_95
#| check: true
gradethis::grade_this({
n <- length(datasets::PlantGrowth$weight)
exp <- qt(0.975, df = n - 1)
x <- .result
if (!is.numeric(x) || length(x)!=1L || !is.finite(x)) fail("Print a single numeric value.")
if (abs(x - exp) < 1e-6) pass("✅ Correct t* for 95%.")
else fail("Check df and make sure it's two-sided 95%.")
})


```

## Q4 — One-Sample Confidence Interval for μ (σ Unknown)

Using PlantGrowth$weight (σ unknown), construct a 95% CI for the true mean plant weight. Give your output as a vector of the form c(lower, upper).

::: {.callout-note title="Preview"}
Run this code chunk to get a glimpse of the dataset. Feel free to change the values to visualize more/less number of rows. 
```{webr}
#| echo: true
data("PlantGrowth")
head(PlantGrowth, 10)
```
:::

::: {.callout-note title="Info"}
When σ is unknown, replace with the sample SD s and use the t distribution:

$$
\mathrm{CI}_{1-\alpha}:\ \bar{x} \pm t_{1-\alpha/2,\;n-1}\,\frac{s}{\sqrt{n}},\ \text{df}=n-1.
$$
:::

```{webr}
#| exercise: q4_t_ci_weight
#| exercise.lines: 1
#| echo: false
data("PlantGrowth")
xbar <- 
sd <- 
n <- 
critical_value <- 


```

::: {.hint exercise="q4_t_ci_weight"}
Form “center ± margin of error”: center is the sample mean; margin uses the t critical value with df = n−1, the sample SD, and √n. Output c(lower, upper).
:::

::: {.solution exercise="q4_t_ci_weight"}
```{webr}
#| exercise: q4_t_ci_weight
#| solution: true
data("PlantGrowth")
xbar <- mean(PlantGrowth$weight); 
sd <- sd(PlantGrowth$weight); 
n <- length(PlantGrowth$weight)
critical_value <- qt(0.975, df = n - 1)
xbar + c(-1,1)*critical_value*sd/sqrt(n)


```
:::

```{webr}
#| exercise: q4_t_ci_weight
#| check: true
gradethis::grade_this({

x <- datasets::PlantGrowth$weight
xbar <- mean(x); s <- sd(x); n <- length(x)
tstar <- qt(0.975, df = n - 1)
exp <- xbar + c(-1,1) * tstar * s / sqrt(n)

r <- .result
ok <- is.numeric(r) && length(r)==2L && all(is.finite(r))
if (!ok) fail("Print a numeric vector c(lower, upper).")
if (max(abs(r - exp)) < 1e-6) pass("✅ Correct 95% t-interval for PlantGrowth$weight.")
else fail("Recheck x̄, s, n, and t* with df = n−1.")
})
```

## Q5 — One sample CI for μ

The dataset used here is river lengths, which gives the lengths (in miles) of 141 “major” rivers in North America, as compiled by the US Geological Survey. Assume river lengths (miles) have known population SD σ = 300. Build a 95% CI for the true mean river length using rivers whe $\alpha = 0.05$.


<figure style="text-align:center">
  <img src="./images/river.jpg" style="width:60%">
  <figcaption>Photo by <a href="https://unsplash.com/@kazuend?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">kazuend</a> on <a href="https://unsplash.com/photos/body-river-surrounded-by-dress-cCthPLHmrzI?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Unsplash</a>
      </figcaption>
</figure>

::: {.callout-note title="Preview"}
Run this code chunk to get a glimpse of the dataset. Feel free to change the values to visualize more/less number of rows. 
```{webr}
#| echo: true
rivers_data <- data.frame(length_miles = as.numeric(datasets::rivers))
head(rivers_data, 10)

```
:::


::: {.callout-note title="Info"}
When σ is unknown, use the t distribution with degrees of freedom n − 1.
:::


```{webr}
#| exercise: q5_zstar_95
#| exercise.lines: 1
#| echo: false
rivers_data <- data.frame(length_miles = as.numeric(datasets::rivers))


```

::: {.hint exercise="q5_zstar_95"}
Find n from the length_miles vector and compute the two-sided 95% z critical value. Then combine it using the formula you have already used in previous questions.
:::

::: {.solution exercise="q5_zstar_95"}
```{webr}
#| exercise: q5_zstar_95
#| solution: true
rivers_data <- data.frame(length_miles = as.numeric(datasets::rivers))
x <- rivers_data$length_miles
xbar <- mean(x); n <- length(x); z <- qnorm(0.975); sigma <- 300
xbar + c(-1,1) * z * sigma / sqrt(n)
```
:::

```{webr}
#| exercise: q5_zstar_95
#| check: true
gradethis::grade_this({
x <- as.numeric(datasets::rivers)
xbar <- mean(x); n <- length(x); z <- qnorm(0.975); sigma <- 300
exp <- xbar + c(-1,1)*z*sigma/sqrt(n)
xres <- .result
ok <- is.numeric(xres) && length(xres)==2L && all(is.finite(xres))
if (!ok) fail("Print a numeric vector c(lower, upper).")
if (max(abs(xres - exp)) < 1e-6) pass("✅ Correct 95% Z-interval with σ known.")
else fail("Recheck x̄, n, z*, and σ = 300.")
})



```


## Q6 — One-Sample CI for μ 

Using airquality$Ozone (ignore missing), construct and print 90% CI for the true
mean ozone level (ppb). Assume normality is reasonable.

<figure style="text-align:center">
  <img src="./images/air_quality.jpg" style="width:60%">
  <figcaption>Photo by <a href="https://unsplash.com/@timwitzdam?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Tim Witzdam</a> on <a href="https://unsplash.com/photos/plant-next-to-a-device-displaying-readings-NH4RJjtDaNk?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Unsplash</a>
      </figcaption>
</figure>


::: {.callout-note title="Preview"}
```{webr}
#| echo: true
aq <- datasets::airquality
oz <- na.omit(aq$Ozone)
head(oz, 10)
```
:::

```{webr}
#| exercise: q6_t_ci_ozone
w <- na.omit(datasets::airquality$Ozone)
xbar <- mean(w)
s <- ＿＿ 
n <- length(w)
tstar <- ＿＿ 
xbar + c(-1,1)*tstar * s/sqrt(n)
```

::: {.hint exercise="q6_t_ci_ozone"}
Two-sided 90% → $t^*$
at 0.95 with df = n−1; use s for the SD.
:::

::: {.solution exercise="q6_t_ci_ozone"}
```{webr}
#| exercise: q6_t_ci_ozone
#| solution: true
w <- na.omit(datasets::airquality$Ozone)
xbar <- mean(w); s <- sd(w); n <- length(w)
tstar <- qt(0.95, df = n - 1) # two-sided 90%
xbar + c(-1,1) * tstar * s / sqrt(n)
```
:::

```{webr}
#| exercise: q6_t_ci_ozone
#| check: true
gradethis::grade_this({
w <- na.omit(datasets::airquality$Ozone)
xbar <- mean(w); s <- sd(w); n <- length(w)
tstar <- qt(0.95, df = n - 1)
exp <- xbar + c(-1,1)*tstar*s/sqrt(n)
r <- .result
ok <- is.numeric(r) && length(r)==2L && all(is.finite(r))
if (!ok) fail("Print a numeric vector c(lower, upper).")
if (max(abs(r - exp)) < 1e-6) pass("✅ Correct 90% t-interval.")
else fail("Check t* at 0.95 and your s, n.")
})
```

## Q7 — One-Sample CI for a Proportion 

In the UCBAdmissions dataset, Admit status for a student is "Admitted" or "Rejected". Build and print c(lower, upper) for the 95% CI on the overall admission proportion.


<figure style="text-align:center">
  <img src="./images/admissions.jpg" style="width:60%">
  <figcaption>Photo by <a href="https://unsplash.com/@mattragland?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Matt Ragland</a> on <a href="https://unsplash.com/photos/flat-lay-photography-of-blue-backpack-beside-book-and-silver-macbook-02z1I7gv4ao?utm_source=unsplash&utm_medium=referral&utm_content=creditCopyText">Unsplash</a>
      </figcaption>
</figure>

::: {.callout-note title="Info"}
For a single proportion:

$$
\mathrm{CI}_{1-\alpha}:\ \hat{p} \pm z_{1-\alpha/2}\,\sqrt{\frac{\hat{p}(1-\hat{p})}{n}}.
$$

Rule: Ensure $n\hat{p}$ and $n(1 - \hat{p})$ are not too small
:::

::: {.callout-note title="Preview"}
```{webr}
#| echo: true
df <- as.data.frame(datasets::UCBAdmissions)  # factors Admit, Gender, Dept + Freq
head(df, 10)
aggregate(Freq ~ Admit, data = df, sum)       # totals by Admit
sum(df$Freq)                                  # total n

```
:::

```{webr}
#| exercise: q7_p_ci_ucb
df <- as.data.frame(datasets::UCBAdmissions)
n <- sum(df$Freq)
phat <- ＿＿  # proportion admitted
z <- ＿＿      
se <- ＿＿ #compute standard error


```

::: {.hint exercise="q7_p_ci_ucb"}
phat <- sum(df$Freq[df$Admit == "Admitted"]) / n and z <- qnorm(0.975) for 95%.
:::

::: {.solution exercise="q7_p_ci_ucb"}
```{webr}
#| exercise: q7_p_ci_ucb
#| solution: true
df <- as.data.frame(datasets::UCBAdmissions)
n <- sum(df$Freq)
phat <- sum(df$Freq[df$Admit == "Admitted"]) / n
z <- qnorm(0.975)
se <- sqrt(phat*(1 - phat)/n)
phat + c(-1,1) * z * se

```
:::

```{webr}
#| exercise: q7_p_ci_ucb
#| check: true
gradethis::grade_this({
df <- as.data.frame(datasets::UCBAdmissions)
n <- sum(df$Freq)
phat <- sum(df$Freq[df$Admit == "Admitted"]) / n
z <- qnorm(0.975)
se <- sqrt(phat*(1 - phat)/n)
exp <- phat + c(-1,1) * z * se

r <- .result
ok <- is.numeric(r) && length(r) == 2L && all(is.finite(r))
if (!ok) fail("Print a numeric vector c(lower, upper).")
if (max(abs(r - exp)) < 1e-6) pass("✅ Correct 95% Wald CI for the admission proportion.")
else fail("Recompute p̂, n, z*, and SE = sqrt(p̂(1−p̂)/n).")
})

```


## Q8 - One-Sample CI for a Variance (σ²)

Using PlantGrowth$weight, print c(lower, upper) for 95% CI on σ^2. (Assume normality.)

::: {.callout-note title="Info"}
For normal data, the variance CI uses the chi-square distribution with df=n−1:

$$
\left(\frac{(n-1)s^2}{\chi^2_{1-\alpha/2,\;n-1}},\; \frac{(n-1)s^2}{\chi^2_{\alpha/2,\;n-1}}\right),\qquad s^2=\text{sample variance}.
$$


For 95%, the cutoffs are at 0.025 and 0.975.
:::

::: {.callout-note title="Preview"}
```{webr}
#| echo: true
data("PlantGrowth")
head(PlantGrowth, 10)
```
:::


```{webr}
#| exercise: q8_var_ci_pg
x <- datasets::PlantGrowth$weight
n <- length(x); df <- ＿＿ ; s2 <- var(x)
chi_lo <- ＿＿ (＿＿, df = df) 
chi_hi <- ＿＿ (0.025, ＿＿)
c(df * s2/chi_lo, df*s2/chi_hi)
```

::: {.hint exercise="q8_var_ci_pg"}
Use $s^2$, df = n−1, and chi-square quantiles at 0.975 and 0.025 (note the order).
:::

::: {.solution exercise="q8_var_ci_pg"}
```{webr}
#| exercise: q8_var_ci_pg
#| solution: true
x <- datasets::PlantGrowth$weight
n <- length(x); df <- n - 1; s2 <- var(x)
chi_lo <- qchisq(0.975, df = df) # upper-tail cutoff
chi_hi <- qchisq(0.025, df = df) # lower-tail cutoff
c(df * s2/chi_lo, df*s2/chi_hi)
```
:::

```{webr}
#| exercise: q8_var_ci_pg
#| check: true
gradethis::grade_this({
x <- datasets::PlantGrowth$weight
n <- length(x); df <- n - 1; s2 <- var(x)
chi_lo <- qchisq(0.975, df = df)
chi_hi <- qchisq(0.025, df = df)
exp <- c(df*s2/chi_lo, df*s2/chi_hi)
r <- .result
ok <- is.numeric(r) && length(r)==2L && all(is.finite(r))
if (!ok) fail("Print a numeric vector c(lower, upper).")
if (max(abs(r - exp)) < 1e-6) pass("✅ Correct 95% CI for σ².")
else fail("Check df=n−1 and chi-square cutoffs (0.025 & 0.975).")
})
```

## Q9 - One-Sample CI for a Proportion (99% Practice)

HairEyeColor is a built-in contingency table of Hair × Eye × Gender of Statistics students with counts. Treat each person as a trial and estimate the true proportion with Brown eyes. Construct a 95% Z-interval for the population proportion p. Print c(lower, upper).

::: {.decor-img}
![](./images/hair_eye_color.jpg){alt="" aria-hidden="true" width="60%" fig-align="center"}
:::

::: {.callout-note title="Info"}
Two-sided 1−α CI for a proportion:

$$
\mathrm{CI}_{1-\alpha}:\ \hat{p} \pm z_{1-\alpha/2}\,\sqrt{\frac{\hat{p}(1-\hat{p})}{n}}.
$$


For 99%, use $z^* = \phi^{-1}$(0.995).
:::

::: {.callout-note title="Preview"}
Here is a glimpse of the UCBAdmissions dataset.
```{webr}
#| echo: true
df <- as.data.frame(datasets::HairEyeColor)
eye_counts <- aggregate(Freq ~ Eye, data = df, sum)
eye_counts
sum(eye_counts$Freq)  # total n

```
:::


```{webr}
#| exercise: q10_z_ci_browneyes
#| exercise.lines: 3
#| echo: false
df <- as.data.frame(datasets::HairEyeColor)
n <- sum(df$＿＿)
phat <- sum(df$Freq[df$＿＿ == "Brown"]) / ＿＿
z <-

```

::: {.hint exercise="q10_z_ci_browneyes"}
Use z <- qnorm(0.975) for 95%.
:::

::: {.solution exercise="q10_z_ci_browneyes"}
```{webr}
#| exercise: q10_z_ci_browneyes
#| solution: true
df <- as.data.frame(datasets::HairEyeColor)
n <- sum(df$Freq)
phat <- sum(df$Freq[df$Eye == "Brown"]) / n
z <- qnorm(0.975)
phat + c(-1,1) * z * sqrt(phat*(1-phat)/n)
```
:::

```{webr}
#| exercise: q10_z_ci_browneyes
#| check: true
gradethis::grade_this({
df <- as.data.frame(datasets::HairEyeColor)
n <- sum(df$Freq)
phat <- sum(df$Freq[df$Eye=="Brown"]) / n
z <- qnorm(0.975)
exp <- phat + c(-1,1)*z*sqrt(phat*(1-phat)/n)

x <- .result
ok <- is.numeric(x) && length(x)==2L && all(is.finite(x))
if (!ok) fail("Print a numeric vector c(lower, upper).")
if (max(abs(x - exp)) < 1e-6) pass("✅ Correct 95% Z-interval for the Brown-eye proportion.")
else fail("Recheck n, phat, and the Z margin for a proportion.")
})

```